# Programming in Python

__Laboratory 02, Type system and built-in types__

__Sept 27, 2022__

__Submission deadline: Sept 29 23:59__

__Submissions must be made through Github Classroom__

Some exercises are graded.
The maximum score you can get is 25 points.
The score value of graded exercises is listed in the instructions.

You are free to use any material (lecture, Stackoverflow etc.) except full solutions but you should first try to figure them out on your own.

We use an [nbgrader](https://nbgrader.readthedocs.io/en/stable/index.html) for autograding the assignments. Many cells are read-only. Please **DO NOT** delete or copy cells. You can add new cells but your solutions should go into the designated cells.

We will use the wikipedia module for some exercises. It can be installed by running pip from Jupyter:

The Python Standard Library has solutions for many of these tasks but you should solve them without relying on the standard library.

We will use the `wikipedia` module in notebook.

```python
import math
! pip install wikipedia
```

Output:
```
Collecting wikipedia
  Using cached wikipedia-1.4.0.tar.gz (27 kB)
  Preparing metadata (setup.py): started
  Preparing metadata (setup.py): finished with status 'done'
Requirement already satisfied: beautifulsoup4 in c:\users\maxig\appdata\local\programs\python\python39\lib\site-packages (from wikipedia) (4.11.1)
Collecting requests<3.0.0,>=2.0.0
  Downloading requests-2.28.1-py3-none-any.whl (62 kB)
Collecting certifi>=2017.4.17
  Downloading certifi-2022.9.24-py3-none-any.whl (161 kB)
Collecting urllib3<1.27,>=1.21.1
  Downloading urllib3-1.26.12-py2.py3-none-any.whl (140 kB)
Collecting idna<4,>=2.5
  Downloading idna-3.4-py3-none-any.whl (61 kB)
Collecting charset-normalizer<3,>=2
  Downloading charset_normalizer-2.1.1-py3-none-any.whl (39 kB)
Requirement already satisfied: soupsieve>1.2 in c:\users\maxig\appdata\local\programs\python\python39\lib\site-packages (from beautifulsoup4->wikipedia) (2.3.2.post1)
Using legacy 'setup.py install' for wikipedia, since package 'wheel' is not
```

# 1. Define a function that takes a sequence as its input and returns whether the sequence is symmetric. A sequence is symmetric if it is equal to its reverse.

<!-- grade_id: cell-9063e2fba9c7157d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def is_symmetric(input):
    return input == input[::-1]
```

<!-- grade_id: cell-58b3ad7e361196e7, points: 0, grade: true, solution: false -->
```python
assert is_symmetric([1]) is True
assert is_symmetric([None, None]) is True
assert is_symmetric([]) is True
assert is_symmetric([1, 2, 3, 1]) is False
assert is_symmetric([1, "foo", "bar", "foo", 1]) is True
assert is_symmetric("abcba") is True
assert is_symmetric(list(range(100)) + list(range(99, -1, -1))) is True
assert is_symmetric(list(range(100)) + list(range(100, 0, -1))) is False
assert is_symmetric([1, "1"]) is False
```

# String manipulation

## 2.1 Define a function that takes a string as its input and return a dictionary with the character frequencies.

Do not use the `collections` module.

<!-- grade_id: cell-7fbaf8610be225a9, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def char_freq(text):
    dict = {}
    for i in text:
        if i in dict:
            dict[i] += 1
        else:
            dict[i] = 1
    return dict
```

<!-- grade_id: cell-ecedb93a6c391539, points: 1, grade: true, solution: false -->
```python
assert char_freq("aba") == {"a": 2, "b": 1}
assert char_freq("") == {}
assert char_freq("ab" * 100 + "cd") == {"a": 100, "b": 100, "c": 1, "d": 1}
```

## 2.2 (1p) Define a function that computes word frequencies in a text.

You can split the text with `text.split()`.

<!-- grade_id: cell-4d4d473dd807f962, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def word_freq(text):
    dict = {}
    t = text.split(" ") if len(text) != 0 else {}
    for i in t:
        if i in dict:
            dict[i] += 1
        else:
            dict[i] = 1
    return dict
```

<!-- grade_id: cell-1c299c2c935129e7, points: 0, grade: true, solution: false -->
```python
s = "the green tea and the black tea"
assert word_freq(s) == {"the": 2, "tea": 2, "green": 1, "black": 1, "and": 1}
assert word_freq("") == {}
```

## 2.3 (2p) Define a function that takes two strings and decides whether they are anagrams. A string is an anagram of another string if its letters can be rearranged so that it equals the other string.¶

For example:

```
abc -- bac
aabb -- abab
```

Counter examples:

```
abc -- aabc
abab -- aaab
```

<!-- grade_id: cell-7ea849b6ffc4421a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def are_anagrams(t1, t2):
    return char_freq(t1) == char_freq(t2)
```

<!-- grade_id: cell-231f78d971a5d48b, points: 2, grade: true, solution: false -->
```python
assert are_anagrams("", "") is True
assert are_anagrams("abc", "bac") is True
assert are_anagrams("aabb", "abab") is True
assert are_anagrams("abab", "aaab") is False
```

## 2.4 (2p) Define a function that takes a string and returns its bigram frequencies as a dictionary.

Character bigrams are pairs of subsequent characters. For example word apple contains the following bigrams: ap, pp, pl, le.

Do not use the `collections` module.

<!-- grade_id: cell-2669edc3d29b28b6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def get_char_bigrams(text):
    dict = {}
    for i in range(len(text) - 1):
        if i == 0:
            dict[text[0] + text[1]] = 1
        else:
            word = text[i] + text[i + 1]
            if word in dict:
                dict[word] += 1
            else:
                dict[word] = 1
    return dict
```

<!-- grade_id: cell-5a3764a1d0909393, points: 2, grade: true, solution: false -->
```python
assert get_char_bigrams("apple") == {"ap": 1, "pp": 1, "pl": 1, "le": 1}
assert get_char_bigrams("apple apple") == {
    "ap": 2,
    "pp": 2,
    "pl": 2,
    "le": 2,
    "e ": 1,
    " a": 1,
}
```

# Wikipedia module

The following exercises use the `wikipedia` package. The basic usage is illustrated below.

The documentation is available [here](https://pypi.python.org/pypi/wikipedia/).

Searching for pages:

```python
import wikipedia

results = wikipedia.search("Vienna")
results
```

Output:
```
['Vienna',
 'Congress of Vienna',
 'Vienna sausage',
 'Battle of Vienna',
 'Vienna Teng',
 'Vienna Award',
 'Austria',
 'Treaty of Vienna',
 'BC Vienna',
 'Vienna Beef']
```

Downloading an article:

```python
article = wikipedia.page("Budapest")

article.summary[:1000]
```

Output:
```
'Budapest (UK: , US: ; Hungarian pronunciation: [ˈbudɑpɛʃt] (listen)) is the capital and most populous city of Hungary. It is the ninth-largest city in the European Union by population within city limits and the second-largest city on the Danube river; the city has an estimated population of 1,752,286 over a land area of about 525 square kilometres (203 square miles). Budapest, which is both a city and county, forms the centre of the Budapest metropolitan area, which has an area of 7,626 square kilometres (2,944 square miles) and a population of 3,303,786, comprising 33% of the population of Hungary.The history of Budapest began when an early Celtic settlement transformed into the Roman town of Aquincum, the capital of Lower Pannonia. The Hungarians arrived in the territory in the late 9th century, but the area was pillaged by the Mongols in 1241–42. Re-established Buda became one of the centres of Renaissance humanist culture by the 15th century. The Battle of Mohács, in 1526, was fol
```

The content attribute contains the full text:

```python
type(article.content), len(article.content)
```

Output:
```
(str, 86778)
```

By default the module downloads the English Wikipedia. The language can be changed the following way:

```python
wikipedia.set_lang("de")
```

```python
wikipedia.search("Budapest")
```

Output:
```
['Budapest',
 'Budapester Memorandum',
 'Grand Budapest Hotel',
 'Schlacht um Budapest',
 'Ungarn',
 'Budapester',
 'Parlamentsgebäude (Budapest)',
 'Metró Budapest',
 'Budapest (Begriffsklärung)',
 'Budapest Keleti pályaudvar']
```

```python
de_article = wikipedia.page("Budapest")
de_article.summary[:100]
```

Output:
```
"Budapest (ungarische Aussprache ['budɒpɛʃt]; ) ist die Hauptstadt und zugleich größte Stadt Ungarns."
```

```python
wikipedia.set_lang("en")
```

## 3.1 (5p) Build a small Wikipedia corpus by downloading pages until their character count exceeds 100000. The corpus should be a list of strings.

You can download arbitrary pages but some pages redirect to disambiguation pages. You should not pick these. Make sure the same page is not included multiple times.

<!-- grade_id: cell-2e7be91b99158dca, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
import random
import string


def build_wikipedia_corpus():
    val = []
    cc = 0
    searching = ""
    while cc <= 100000:
        searching += random.choice(string.ascii_letters)
        search_ = wikipedia.search(searching)
        for i in search_:
            try:
                cur = wikipedia.page(i, auto_suggest=False).content
                val.append(cur)
                cc += len(cur)
            except wikipedia.exceptions.DisambiguationError as e:
                searching = random.choice(string.ascii_letters)
    return val
```

<!-- grade_id: cell-f26aa917928c62c8, points: 5, grade: true, solution: false -->
```python
corpus = build_wikipedia_corpus()
# check types
assert isinstance(corpus, list)
assert isinstance(corpus[0], str)

# check character count
sum_length = 0
for document in corpus:
    sum_length += len(document)
assert sum_length >= 100000

# no duplicate pages
assert len(set(corpus)) == len(corpus)
```

Output:
```
C:\Users\maxig\AppData\Local\Programs\Python\Python39\lib\site-packages\wikipedia\wikipedia.py:389: GuessedAtParserWarning: No parser was explicitly specified, so I'm using the best available HTML parser for this system ("lxml"). This usually isn't a problem, but if you run this code on another system, or in a different virtual environment, it may use a different parser and behave differently.

The code that caused this warning is on line 389 of the file C:\Users\maxig\AppData\Local\Programs\Python\Python39\lib\site-packages\wikipedia\wikipedia.py. To get rid of this warning, pass the additional argument 'features="lxml"' to the BeautifulSoup constructor.

  lis = BeautifulSoup(html).find_all('li')
```

## 3.2 (2p) Add a `min_char` parameter to `build_wikipedia_corpus` that defaults to 10000. You should stop building the corpus when this number is reached.

<!-- grade_id: cell-bf7ffeba0644e753, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def build_wikipedia_corpus(min_char=10000):
    val = []
    cc = 0
    searching = ""
    while cc < min_char:
        searching += random.choice(string.ascii_letters)
        search_ = wikipedia.search(searching)
        for i in search_:
            try:
                cur = wikipedia.page(i, auto_suggest=False).content
                if cc + len(cur) < min_char:
                    val.append(cur)
                    cc += len(cur)
                else:
                    val.append(cur[:min_char - cc])
                    cc = min_char
                    break
            except wikipedia.exceptions.DisambiguationError as e:
                searching = random.choice(string.ascii_letters)
    return val
```

<!-- grade_id: cell-68c3de00731699f8, points: 2, grade: true, solution: false -->
```python
corpus = build_wikipedia_corpus(1000)
# check types
assert isinstance(corpus, list)
assert isinstance(corpus[0], str)

# check character count
sum_length = 0
for document in corpus:
    sum_length += len(document)
assert sum_length >= 1000
# shouldn't be too long
assert sum_length < 100000
```

## 3.3 (3p) Compute the word frequencies of your Wikipedia corpus. Return the $k$ most frequent words with their frequencies as a list of tuples. $k$ should default to 10. The list should be ordered by frequency starting from the most frequent.

You can use str.split() for word tokenization.

<!-- grade_id: cell-8c167ceb2a4c65f1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def get_corpus_word_frequencies(corpus, k =10):
    d = {}
    for i in corpus:
        d = d | word_freq(i)
    return sorted(d.items(), key=lambda x: x[1], reverse=True)[:k]
```

<!-- grade_id: cell-d17a9e499c13bd35, points: 3, grade: true, solution: false -->
```python
top_words_and_freqs = get_corpus_word_frequencies(corpus)
assert isinstance(top_words_and_freqs, list)
assert len(top_words_and_freqs) == 10
assert isinstance(top_words_and_freqs[0], tuple)

# Frequency of the most frequent word
fr = top_words_and_freqs[0][1]

for word, freq in top_words_and_freqs:
    assert isinstance(word, str)
    assert isinstance(freq, int)
    # Check decreasing order
    assert freq <= fr
    fr = freq

top_words_and_freqs5 = get_corpus_word_frequencies(corpus, k=5)
assert len(top_words_and_freqs5) == 5
```

## 3.4.2 Print the ten most frequent words and their frequency in the following format:

```
a<TAB>12
an<TAB>11
```

<!-- grade_id: cell-bc2d984be320ff7c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In [128], line 2
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

# Linear algebra

## 4.1 (2p) Define a function that takes a list of lists and decided if it is a matrix. A list of lists is a matrix if each nested list is the same length.

<!-- grade_id: cell-90b341f4b426d33d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def is_matrix(list):
    if len(list) > 1:
        length = len(list[0])
        for i in list[1:]:
            if length != len(i):
                return False
    return True
```

<!-- grade_id: cell-218d090c410748b5, points: 2, grade: true, solution: false -->
```python
assert is_matrix([]) is True
assert is_matrix([[], []]) is True
assert is_matrix([[1, 2], [2, 1]]) is True
assert is_matrix([[1], [2, 1]]) is False
assert is_matrix([[1, 2], [2, 1], [3, 4]]) is True
assert is_matrix([[1, 2], [2, 1], [3, 4, 5]]) is False
```

## 4.2 (2p) Define a function that takes a matrix as an input represented as a list of lists. Return its transpose without changing the original matrix.

Raise a `ValueError` if the input is not a valid matrix. Don't forget to reuse your code.

<!-- grade_id: cell-adee24575dcbac8e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
import numpy as np


def transpose(matrix):
    if is_matrix(matrix):
        return np.transpose(matrix).tolist()
    else:
        raise ValueError
```

<!-- grade_id: cell-9b412cb3d0c0506a, points: 2, grade: true, solution: false -->
```python
m1 = [[1, 2, 3], [4, 5, 6]]
m2 = [[1, 4], [2, 5], [3, 6]]

assert transpose([[1]]) == [[1]]
assert transpose(m1) == m2
# m1 is unchanged
assert m1 == [[1, 2, 3], [4, 5, 6]]

# transpose is an involution (self-inverse)
assert transpose(transpose(m1)) == m1
```

## 4.3 (1p) Define a function that takes a matrix as its input and returns the row-wise maximum elements as a list.

Raise a `ValueError` if the input is not a valid matrix. Don't forget to reuse your code.

<!-- grade_id: cell-21f91aada57c0332, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def rowwise_max(matrix):
    if is_matrix(matrix):
        return np.amax(matrix, axis=1).tolist()
    else:
        raise ValueError
```

<!-- grade_id: cell-eb440f91ea977032, points: 1, grade: true, solution: false -->
```python
assert rowwise_max(
    [
        [1, 2, 3],
        [1, 4, 3],
        [1, 5, 3],
    ]
) == [3, 4, 5]

assert rowwise_max([list(range(1000))]) == [999]
assert rowwise_max(transpose([list(range(1000))])) == list(range(1000))
```

## 4.4 (2p) Define a function that takes a matrix as its input and returns the column-wise maximum elements as a list.

Try to reuse your previous functions.

<!-- grade_id: cell-bde0cb9cf8d75c89, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def columnwise_max(matrix):
    return np.amax(matrix, axis=0).tolist()
```

<!-- grade_id: cell-142e6034ed15005b, points: 2, grade: true, solution: false -->
```python
assert columnwise_max(
    [
        [1, 2, 3],
        [1, 4, 3],
        [1, 5, 3],
    ]
) == [1, 5, 3]

assert columnwise_max([list(range(1000))]) == list(range(1000))
assert columnwise_max(transpose([list(range(1000))])) == [999]
```

## 4.5 (3p) Define a function that takes two vectors represented as lists and returns their L2 distance.

Also called Euclidean distance.

<!-- grade_id: cell-849b6ea4c65f058e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def l2_distance(m1, m2):
    return np.linalg.norm(np.array(m1) - np.array(m2))
```

<!-- grade_id: cell-4c9f06a8534cbefc, points: 3, grade: true, solution: false -->
```python
assert abs(l2_distance([1, 0], [1, 0]) - 0) < 1e-5
assert abs(l2_distance([1, 0], [0, 1]) - 2 ** 0.5) < 1e-5
```

## 4.6 Create a function that takes a matrix of size NxM and returns the pairwise L2 distance of each row pair.

For input $A^{N \times M}$, the output should be $D^{N \times N}$ where: $D_{i,j} = L_2 (A_i, A_j)$

<!-- grade_id: cell-644d3391fc995ee9, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def pairwise_row_distance(matrix):
    d = []
    for idx,row in enumerate(matrix):
        d.append([])
        for col in matrix:
            d[idx].append(l2_distance(row,col))
    return d
```

<!-- grade_id: cell-8bc3cf03b303c7c3, points: 0, grade: true, solution: false -->
```python
D = pairwise_row_distance([[1, 0, 2], [0, 2, 1]])
# should be a matrix
assert is_matrix(D) is True
assert len(D) == 2
assert len(D[0]) == 2
assert D[0][0] == 0
assert abs(D[0][1] - 2.449489742783178) < 1e-3
assert abs(D[1][0] - 2.449489742783178) < 1e-3
assert D[1][1] == 0

# should be symmetric
assert transpose(D) == D
```