Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Max"
COLLABORATORS = "/"
```

---

# __Assignment 3: Object oriented programming__

Imports

```python
from typing import Any, Dict, List
```

## 1. Basic class definition

### 1.1  Define a class named `A` with an `__init__` function that takes a single parameter and stores it in an attribute named `value`. Add a `print_value` method to the class, that prints out the value set by the `__init__` function.

Instantiate the class and call the `print_value` method. Please use type hints, and the value parameter can be Any type.

<!-- grade_id: cell-6c7ba38d59f0b533, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
class A:
    def __init__(self, value):
        self.value = value

    def print_value(self):
        print(self.value)
```

```python
a = A("abc")
a.print_value()
```

Output:
```
abc
```

### 1.2 Redefine the class's `__init__` so that it can be instantiated without a parameter. If it is called without a parameter, value should be 42.

<!-- grade_id: cell-1fa4f4a04cc3b613, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
class A:
    def __init__(self, value="42"):
        self.value = value

    def print_value(self):
        print(self.value)
```

```python
a = A("abc")
a.print_value()  # prints abc
a = A()
a.print_value()  # prints 42
```

Output:
```
abc
42
```

### 1.3 Define a class named `B`, whose `__init__` takes two parameters and stores one in a public attribute and the other in a private attribute named `this_is_public` and `__this_is_private` respectively.

Check the class's __dict__ attribute and find out the mangled name of the private attribute. Don't forget the type hints!

<!-- grade_id: cell-f4c8b3c52d081013, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
class B:
    def __init__(self, public, private):
        self.this_is_public = public
        self.__this_is_private = private
```

<!-- grade_id: cell-6480bb6183af4ff0, points: 0, grade: true, solution: false -->
```python
b = B(1, 2)
assert b.this_is_public == 1
try:
    b.__this_is_private
    print("This should not happen.")
except AttributeError:
    print("Failed to access private attribute, this is good :)")
```

Output:
```
Failed to access private attribute, this is good :)
```

## 2. Inheritance

### 2.1 Guess the output without running the cell.

```python
class A(object): pass


class B(A): pass


class C(A): pass


class D(B): pass


a = A()
b = B()
c = C()
d = D()

print(isinstance(a, object))  # True
print(isinstance(b, object))  # True
print(isinstance(a, B))  # False
print(isinstance(b, A))  # True
print(isinstance(d, A))  # True
```

Output:
```
True
True
False
True
True
```

```python
print(issubclass(C, object))  # True
print(issubclass(D, B))  # True
print(issubclass(B, D))  # False
print(issubclass(B, B))  # True
```

Output:
```
True
True
False
True
```

### 2.2 Create a Cat, a Dog, a Fish and a Eagle class.

The animals have the following attributes:

1. cats, dogs and eagles can make a sound (this should be a make_sound function that prints the animals sound),
2. all animals have an age and a number_of_legs attribute,
3. cats and dogs have a fur_color attribute. They can be instantiated with a single color or a list or tuple of colors.

Use inheritance and avoid repeating code. Use default values in the constructors.

<!-- grade_id: cell-45b3470b12151035, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
class Animal:
    def __init__(self, name, age=None):
        self.name = name
        self.age = age


class Cat(Animal):
    def __init__(self, name, age, fur):
        super().__init__(name, age)
        self.fur = fur

    @staticmethod
    def make_sound():
        print("Miau")


class Dog(Cat):
    def __init__(self, name, age, fur):
        super().__init__(name, age, fur)
        self.sound = "Wuff"

    @staticmethod
    def make_sound():
        print("Wuff")


class Fish(Animal):
    @staticmethod
    def make_sound():
        pass


class Eagle(Animal):
    def __init__(self, name, age):
        super().__init__(name, age)

    @staticmethod
    def make_sound():
        print("Qia")

```

<!-- grade_id: cell-b584e9c19590038b, points: 0, grade: true, solution: false -->
```python
try:
    cat = Cat("Fluffy", age=3, fur="white")
    dog = Dog("Cherry", age=1, fur=("white", "brown", "black"))
    fish = Fish("King")
    eagle = Eagle("Bruce", age=2)
    animals = [cat, dog, fish, eagle]
except Exception as e:
    print(f'Could not initialize an animal: {e}')
```

Iterate over the list animals and call `make_sound` for each animal. Print either the sound the animal makes or "XY does not make a sound" if the animal does not make a sound (fish). This is an example of duck typing.

```python
for animal in animals:
    animal.make_sound()
```

Output:
```
Miau
Wuff
Qia
```

## ================ GRADED PART ================

## 3. Graded exercises

### 3.1 Cookbook [sum 15 points]

#### 3.1.1 Recipe class [5 points]
Create a new class called Recipe. A recipe should have a name (string), a dictionary of ingredients (name of the ingredient and the amount of it) and a list of steps (string instructions). Define a `list_ingredients` function, that returns a string of ingredients, with one line per ingredient and the ingredient and the amount are tabulator separated format. Define a `list_steps` function that returns a string of the steps as a numbered list.

Additionally make it so that the `len()` function called on a Recipe will return the number of ingredients.

<!-- grade_id: cell-f51366c897f51556, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
class Recipe:
    def __init__(self, name, ingredients, steps):
        self.name = name
        self.ingredients = ingredients
        self.steps = steps

    def list_ingredients(self):
        _str = ""
        for ingredient, amount in zip(self.ingredients.keys(), self.ingredients.values()):
            _str += (ingredient + "\t" + str(amount)) + "\n"
        return _str

    def list_steps(self):
        _str = ""
        for idx, task in enumerate(self.steps):
            _str += str(idx + 1) + ". " + task + "\n"
        return _str

    def __len__(self):
        return len(self.ingredients.keys())
```

```python
ingredients = {"flour": 1, "egg": 2}
steps = ["Mix the ingredients well", "Roll the dough thinly"]
pasta = Recipe("pasta", ingredients, steps)
```

The output should look like this:

pasta.list_ingredients()

```
flour   1
egg 2
```

pasta.list_steps()

```
1. Mix the ingredients well
2. Roll the dough thinly
```

<!-- grade_id: cell-607ca863bc22b530, points: 5, grade: true, solution: false -->
```python
assert 'flour\t1\negg\t2' in pasta.list_ingredients()
assert pasta.list_steps().startswith("1.")
assert len(pasta) == 2
```

#### 3.1.2 ComplexRecipe class [7 points]

Define a new class based on the previous Recipe class, that can also take a list of recipes as an optional parameter (by default it should be None). If the recipes attribute is not None when the `list_ingredients` or the `list_steps` it should start the output string with the the name and ingredients/steps of the recipes in the recipe list and then after a new line, print out it's own ingredients/steps.

Redifine the lenght of the object, so it includes the lengths of the included recipes as well.

<!-- grade_id: cell-71160792c789ff85, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
class ComplexRecipe(Recipe):
    def __init__(self, name, ingredients, steps, recipes=None):
        super().__init__(name, ingredients, steps)
        self.recipes = recipes

    def list_ingredients(self):
        _str = ""
        if self.recipes is not None:
            for i in self.recipes:
                _str += i.name + "\n"
                for ingredient, amount in zip(i.ingredients.keys(), i.ingredients.values()):
                    _str += (ingredient + "\t" + str(amount)) + "\n"
            _str += "\n"
        for ingredient, amount in zip(self.ingredients.keys(), self.ingredients.values()):
            _str += (ingredient + "\t" + str(amount)) + "\n"
        return _str

    def list_steps(self):
        _str = ""
        if self.recipes is not None:
            for i in self.recipes:
                _str += i.name + "\n"
                for idx, task in enumerate(i.steps):
                    _str += str(idx + 1) + ". " + task + "\n"
            _str += "\n"
        for idx, task in enumerate(self.steps):
            _str += str(idx + 1) + ". " + task + "\n"
        return _str

    def __len__(self):
        _steps = len(self.ingredients.keys())
        for i in self.recipes:
            _steps += len(i.ingredients.keys())
        return _steps
```

```python
ragout_ingredients = {"meat": 1, "celery": 2, "carot": 3, "tomato paste": 1, "spices": 2}
ragout_steps = ["Cook the vegetables with spices", "Add the meat and let it brown", "Add tomato paste and cook."]
ragout = Recipe("ragout", ragout_ingredients, ragout_steps)
besamel_ingredients = {"flour": 3, "butter": 3, "milk": 5, "nutmeg": 1}
besamel_steps = ["Melt the butter", "Add flour and nutmeg", "Slowly stir in the milk until incorporated"]
besamel = Recipe("besamel", besamel_ingredients, besamel_steps)

lasagne_ingredients = {"parmesan": 3}
lasagne_steps = ["Layer ragout, besamel, and cooked pasta", "Top it off with parmesan", "Cook in the oven"]
lasagne = ComplexRecipe("lasagne", lasagne_ingredients, lasagne_steps, recipes=[pasta, ragout, besamel])
```

The output should look like this:

lasagne.list_ingredients()

```
pasta
flour	1
egg	2
ragout
meat	1
celery	2
carot	3
tomato paste	1
spices	2
besamel
flour	3
butter	3
milk	5
nutmeg	1

parmesan	3
```

lasagne.list_steps()

```
pasta
1. Mix the ingredients well
2. Roll the dough thinly
ragout
1. Cook the vegetables with spices
2. Add the meat and let it brown
3. Add tomato paste and cook.
besamel
1. Melt the butter
2. Add flour and nutmeg
3. Slowly stir in the milk until incorporated

1. Layer ragout, besamel, and cooked pasta
2. Top it off with parmesan
3. Cook in the oven
```

<!-- grade_id: cell-2ef9ed4bc1003b26, points: 7, grade: true, solution: false -->
```python
assert "pasta" in lasagne.list_ingredients()
assert lasagne.list_steps().count("1") == 4
assert len(lasagne) == 12
```

#### 3.1.3 Cookbook class [3 points]

The Cookbook class should have an author, a title and a list of recipes. Make it so that the whole book can be printed in a nice format by overriding the `__str__` function. The outpus should contain the name of the author, the title of the book, and the recipes using the previously defined functions, or adding a `__str__` to those classes as well.

<!-- grade_id: cell-5e027ff98851d395, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
class Cookbook:
    def __init__(self, author, title, recipes):
        self.author = author
        self.title = title
        self.recipes = recipes

    def __str__(self):
        _str = self.author + "\n" + self.title + "\n"
        for i in self.recipes:
            _str += i.name + "\n" + i.list_ingredients() + "\n" + i.list_steps() + "\n\n"
        return _str
```

```python
cookbook = Cookbook("John Doe", "Big Lasagne Book", [pasta, lasagne])
```

The output can be something like this:
```
John Doe
Big Lasagne Book
pasta
flour	1
egg	2

1. Mix the ingredients well
2. Roll the dough thinly


lasagne
pasta
flour	1
egg	2
ragout
meat	1
celery	2
carot	3
tomato paste	1
spices	2
besamel
flour	3
butter	3
milk	5
nutmeg	1

parmesan	3

pasta
1. Mix the ingredients well
2. Roll the dough thinly
ragout
1. Cook the vegetables with spices
2. Add the meat and let it brown
3. Add tomato paste and cook.
besamel
1. Melt the butter
2. Add flour and nutmeg
3. Slowly stir in the milk until incorporated

1. Layer ragout, besamel, and cooked pasta
2. Top it off with parmesan
3. Cook in the oven
```

<!-- grade_id: cell-4ead2233848edb80, points: 3, grade: true, solution: false -->
```python
assert str(cookbook).count("pasta") == 4
```

### 3.2 Exceptions [sum 10 points]

#### 3.2.1 Negative Exception [3 points]

Write a `NegativeException` class that is derived from the Exception! Write a `simple_function` that takes an int argument and raises a NegativeException if the argument was negative and returns the argument otherwise.

<!-- grade_id: cell-82844608c165f0ee, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
class NegativeException(Exception):
    pass


def simple_function(number):
    if number < 0:
        raise NegativeException
    return number
```

<!-- grade_id: cell-f1f6b331e2ebf81e, points: 3, grade: true, solution: false -->
```python
assert simple_function(2) == 2

raised = False
try:
    simple_function(-2)
except NegativeException:
    raised = True

assert raised
```

#### 3.2.2 Cascading exceptions [4 points]

Write a `cascade_function` that takes one argument with cascading try-except block that calls the `simple_function` and raises a `ValueError` if a `NegativeException` occured.

<!-- grade_id: cell-8a7b97c572177d2a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def cascade_function(number):
    try:
        return simple_function(number)
    except NegativeException:
        raise ValueError
```

<!-- grade_id: cell-809ca2a34720dfa8, points: 4, grade: true, solution: false -->
```python
negative_raised = False
value_raised = False
try:
    cascade_function(-2)
except NegativeException:
    negative_raised = True
except ValueError:
    value_raised = True

assert not negative_raised
assert value_raised
```

#### 3.2.3 Finally [3 points]

Write a `finally_function` that takes an argument, calls the `cascade_function`, catches any error and as a clean-up action it returns True.

<!-- grade_id: cell-fe8bae55d49a6ec0, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def finally_function(number):
    try:
        cascade_function(number)
    except Exception:
        return True
```

<!-- grade_id: cell-bb0801a5c48f99ac, points: 3, grade: true, solution: false -->
```python
assert finally_function(-5)
```